ATP-driven conformational dynamics reveal hidden intermediates in a heterodimeric ABC transporter

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

Summary:

Pecak et al have deciphered the conformational dynamics of a heterodimeric model ABC transporter, TmrAB, a functional homolog of the human antigen transporter TAP, using single molecule Forster resonance energy and fluorophores attached to residues at either nucleotide binding domains or periplasmic gate. The analysis not only differentiated ATP-free and bound states, but also enabled the real time monitoring of protein conformational changes precisely dissecting transport cycles and resolving transient intermediates. This study is absolutely significant in providing and establishing a general pipeline delineating the conformational dynamics in heterodimeric ABC transporters.

Strengths:

The scientific study is very well documented for experimental design, results and conclusions supported by the experimental data. Authors have determined the conformational dynamics of TmrAB across different ATP concentrations including physiological ones and resolved an outward open state and other conformational states consistent with previous cryoEM and DEER studies. Authors have also mentioned limitations in the study.

Comments on revised version.

Authors have worked on most of the revisions stated in previous feedback and included in the newer version, which has been significantly improved. Other comments have been described to be out of scope from this study.

Reviewer #2 (Public review):

In their manuscript entitled 'ATP-driven conformational dynamics reveal hidden intermediates in a heterodimeric ABC transporter', Pečak et al. use elegant single-molecule FRET experiments in detergent to investigate the heterodimeric ABC transporter TmrAB. By combining simulations of the transporter's accessible volume with elegant trapping strategies, the authors identify an unresolved outward-facing open state and conclude that it is usually obscured by a rapidly interconverting ATPbound ensemble. Overall, the study demonstrates that smFRET can resolve the short-lived intermediate states of TmrAB and potentially other ABC transporters that are obscured in ensemble measurements.

It is a very interesting study that highlights the power of combining high-resolution structural information with spectroscopic approaches. I had three major concerns with the original version, all of which have been addressed by the authors in this revised version.

Recommendations for the authors:

Reviewer #2 (Recommendations for the authors):

I mentioned that the final section of the Results part seems like an afterthought, especially since the heading suggests a broader scope.

Reply: We appreciate this comment. We have revised the final section of the Results to improve its structure and ensure that the scope indicated by the heading is fully reflected in the content. This section now more clearly integrates kinetic and thermodynamic aspects of the transport cycle.

The changes made to the section do not align with the wording of the reply. Please consider modifying it further.

We appreciate the positive feedback and this final comment. We have revised the final section of the Results to better reflect the scope indicated by the heading. In addition to clarifying the kinetic analysis, we now explicitly relate our kinetic observations to previously determined thermodynamic measurements, showing that the rapid interconversion of ATP-bound conformations observed during steady-state turnover is consistent with a thermodynamic landscape characterized by a near-zero free-energy difference and entropy–enthalpy compensation. This revision more clearly integrates the kinetic and thermodynamic aspects of the transporter cycle.

Changes from the earlier version to the revised version

Abstract

ATP-binding cassette (ABC) transporters are essential molecular machines whose conformational dynamics have largely been inferred from ensemble-averaged measurements. Resolving dynamic heterogeneity and transient intermediates, however, requires single-molecule approaches. Here, we use single-molecule Förster resonance energy transfer (smFRET) to resolve ATP-driven conformational dynamics of the heterodimeric type IV ABC transporter TmrAB, a functional homolog of the human antigen transporter TAP, at the level of individual molecules. Fluorophores positioned at the nucleotide-binding domains and periplasmic gate were validated by accessible-volume simulations, fluorescence lifetimes, and ensemble FRET, demonstrating that these reporters reliably track conformational transitions. Single-molecule analysis distinguishes ATP-free and ATP-bound states and quantifies ATP-dependent population shifts from nucleotide-free to physiological ATP concentrations. Kinetic analysis further reveals an unexpectedly long ATP-bound dwell time of ∼300 ms. Using complementary stabilization strategies, we directly resolve a previously hidden outward-facing open state that is kinetically masked under turnover conditions. These results provide the first single-molecule characterization of TmrAB and establish a quantitative single-molecule framework for dissecting ATP-coupled conformational dynamics in heterodimeric ABC transporters.

Introduction

ATP-binding cassette (ABC) transporters constitute the largest family of primary active membrane transport systems, conserved across all domains of life1–3. Despite considerable structural diversity, all ABC transporters share a modular architecture comprising two conserved nucleotide-binding domains (NBDs)–the defining hallmark of the family–and two transmembrane domains (TMDs) that form the substrate translocation pathway3,4. Based on their TMD architecture, ABC transporters are classified into seven types that encompass importers, exporters, extractors, and mechanotransmitters5. Substrate translocation is driven by large conformational changes that are chemo-mechanically coupled to ATP binding, hydrolysis, and phosphate/ADP release2,3. ABC transporters play central roles in cellular homeostasis, nutrient uptake, waste removal, and toxin defense. Their dysfunction and misregulation are linked to numerous diseases and drug resistance6.

The heterodimeric type IV ABC transporter TmrAB from Thermus thermophilus has emerged as a powerful model system due to its exceptional thermal stability and functional homology to the transporter associated with antigen processing (TAP1/2), a key component of adaptive immunity7–9. Notably, TmrAB shares overlapping peptide specificity with TAP and can restore antigen presentation in TAP-deficient human cells10. Its inherent asymmetry, with one catalytically active (canonical) and one inactive (noncanonical) nucleotide-binding site (NBS), provides a unique opportunity to investigate functional specialization and asymmetry in ABC transport mechanisms.

Extensive structural studies, particularly using cryogenic electron microscopy (cryo-EM), have delineated the conformational landscape of TmrAB and yielded a detailed model of its translocation cycle11,12. In this model, TmrAB fluctuates between inward-facing wide and narrow conformations (IFwide and IFnarrow), characterized by a sealed periplasmic gate (PG) and well-separated NBDs, thereby permitting substrate access to the central binding cavity. ATP binding to both NBDs induces NBD dimerization and drives the transition into the outward-facing (OF) states, including an OF open (OFopen) conformation with an open PG that enables substrate release into the periplasm, as well as an OF occluded (OFoccluded) state characterized by a sealed PG and dimerized NBDs. Subsequent ATP hydrolysis and phosphate release lead to asymmetric unlocked return states (URasym and URasym*), before the transporter returns to the IF conformation. These UR states feature a sealed PG, a partially open ADP-bound canonical NBS, and a tightly ATP-occluded noncanonical NBS11.

Single-turnover experiments have shown that ATP binding, rather than hydrolysis, drives the IF-to-OF transition, while phosphate release precedes the OF-to-IF switch12–14. Complementary ensemble approaches, including pulsed electron–electron double resonance (PELDOR/DEER) spectroscopy, have further characterized ATP-dependent conformational changes15,16. However, ensemble averaging inherently masks molecular heterogeneity, obscures inactive or misfolded subpopulations, and limits access to kinetic information.

Single-molecule techniques overcome these limitations by resolving conformational dynamics at the level of individual molecules17,18. In particular, single-molecule Förster resonance energy transfer (smFRET) enables real-time monitoring of protein conformational changes with nanometer precision19–22. Applied to ABC transporters, smFRET provides a unique opportunity to dissect transport cycles, resolve transient intermediates, and extract kinetic and mechanistic insights that remain inaccessible to ensemble-based measurement approaches23–25.

Earlier version: Here, we apply total internal reflection fluorescence (TIRF) microscopy combined with alternating laser excitation (ALEX)-based smFRET to detergent-solubilized heterodimeric ABC transporter TmrAB, providing the first single-molecule characterization of this system. By strategically positioning fluorophore pairs, we directly monitor ATP-dependent NBD dimerization and periplasmic gate (PG) opening, quantify conformational state occupancies across ATP concentrations ranging from nucleotide-free to physiological levels (3 mM), and uncover conformational dynamics previously masked by ensemble averaging. Using three orthogonal trapping strategies–(i) a slow-turnover mutant11,12, (ii) Mg2+ depletion14,15, and (iii) substrate trans-inhibition26,27–we resolved a previously hidden outward-facing open (OFopen) state that rapidly exchanges with the outward-facing occluded (OFoccluded) state. Distance measurements derived from smFRET closely matched predictions from accessible-volume (AV) simulations, cryo-EM structures, and PELDOR/DEER spectroscopy, confirming that detergent-solubilized TmrAB retains a native-like conformational landscape. Together, these results provide the first single-molecule quantification of conformational state occupancies for a heterodimeric type IV ABC transporter and establish TmrAB as a versatile model for single-molecule studies of ABC transport systems.

Revised version: Here, we apply total internal reflection fluorescence (TIRF) microscopy combined with alternating laser excitation (ALEX)-based smFRET to detergent-solubilized heterodimeric ABC transporter TmrAB, providing the first single-molecule characterization of this system. By strategically positioning fluorophore pairs, we directly monitor ATP-dependent NBD dimerization and periplasmic gate (PG) opening, quantify conformational state occupancies across ATP concentrations ranging from nucleotide-free to physiological levels (3 mM), and uncover conformational dynamics previously masked by ensemble averaging. Using three orthogonal trapping strategies–(i) a slow-turnover mutant11,12, (ii) Mg²⁺ depletion14,15, and (iii) substrate trans-inhibition26,27–we resolved a previously hidden outward-facing open (OFopen) state that rapidly exchanges with the outward-facing occluded (OFoccluded) state. Distance measurements derived from smFRET closely matched predictions from accessible-volume (AV) simulations, cryo-EM structures, and PELDOR/DEER spectroscopy, confirming that detergent-solubilized TmrAB retains a native-like conformational landscape. Together, these results provide the first single-molecule quantification of conformational state occupancies for a heterodimeric type IV ABC transporter and establish TmrAB as a versatile model for single-molecule studies of ABC transport systems.

Results

Design of FRET-labeled TmrAB variants to probe conformational dynamics

To monitor conformational changes in distinct regions of TmrAB, we engineered double-cysteine variants for smFRET targeting the nucleotide-binding domains (NBDs) and the periplasmic gate (PG). The NBDs undergo ATP-dependent dimerization and post-hydrolysis dissociation, whereas the PG opening and closing controls substrate release into the periplasm2,3,11. Probing both regions provides complementary readouts of cytosolic and periplasmic coupling during transport.

Labeling positions were selected based on prior PELDOR/DEER studies16. The NBD reporter variant (TmrAC416BL458C, referred to as TmrABNBD) monitors the noncanonical nucleotide-binding site (NBS), while the PG reporter (TmrAC416A, T61CBR56C, hereafter TmrABPG) reports PG opening. In TmrABNBD, the native single cysteine (C416) was retained for labeling, whereas in TmrABPG it was substituted by alanine to prevent off-target labeling. Using the noncanonical NBS prevents direct distinction between OFoccluded and asymmetric unlocked return states (URasym and URasym*)11, reducing the number of resolvable FRET states and simplifying data interpretation.

Both variants were labeled with photostable LD555/LD655 fluorophores containing a 1,3,5,7-cyclooctatetraene moiety to suppress photobleaching and blinking28,29. Accessible-volume (AV) simulations30 across nine cryo-EM structures11 confirmed that donor–acceptor distances (RDA) and simulated FRET efficiencies (Esim) lie within the measurable range (Fig. 1 and Table 1). For TmrABNBD, Esim shifts from 0.62 ± 0.02 (57.9 ± 0.7 Å, NBDs separated) to 0.84 ± 0.01 (45.2 ± 0.4 Å, NBDs dimerized), corresponding to predicted ΔEsim of 0.22 ± 0.02 and ΔRDA of 12.7 ± 0.8 Å. For TmrABPG, Esim changes from 0.96 ± 0.02 (30.6 ± 0.4 Å, closed PG) to 0.69 ± 0.03 (53.8 ± 2.0 Å, open PG), corresponding to predicted ΔEsim of 0.27 ± 0.04 and ΔRDA of 23.2 ± 2.0 Å.

Accessible-volume (AV) simulations of LD fluorophores on TmrAB variants.

AV simulations were performed for LD555 (donor) and LD655 (acceptor) fluorophores attached to the selected TmrAB labeling sites to assess whether donor–acceptor distances are suitable for smFRET measurements30. a, TmrABNBD (TmrAC416BL458C) and b, TmrABPG (TmrAC416A, T61CBR56C) in the inward-facing wide (IFwide; PDB: 6RAN, left) and outward-facing open (OFopen; PDB: 6RAH, right) conformations. The approximate membrane position is indicated by the dashed grey line. TmrA is shown in blue with LD655 (orange) and TmrB in yellow with LD555 (green), with labeling positions indicated (416/458 for TmrABNBD; 61/56 for TmrABPG). For illustration, fluorophores are shown in a defined configuration; however, in experiments donor and acceptor dyes are stochastically attached, resulting in random distribution between subunits.

A slow-turnover TmrAB variant (TmrAC416A, E523Q, T61CBR56C, hereafter TmrABPG_EQ) was also generated to kinetically stabilize ATP-bound conformations. This mutation reduces ATP hydrolysis by ∼1000-fold, extending the catalytic half-life of ∼25 min at 45 °C11,12,14,15.

TmrAB variants are suitable for FRET studies

TmrAB variants were expressed in E. coli and purified by immobilized metal-affinity chromatography. SDS-PAGE and size-exclusion chromatography (SEC) confirmed purity and monodispersity of the samples (Fig. 1–Fig. S1a,b). ATPase assays showed that TmrABwt retained enzymatic activity after purification, yielding a Michaelis-Menten constant (Km) of 0.97 ± 0.28 mM and turnover rate (kcat) of 2.57 ± 0.38 s-1 at 40 °C (Fig. 1–Fig. S1c), consistent with previous reports9,10,12.

Site-specific labeling yielded >90% labeling efficiency per cysteine, with ∼40–50% per-site labeling efficiency for donor-only and acceptor-only populations (Fig. 1–Fig. S1d–f). Stochastic labeling of the TmrAB variants results in both homo (donor-donor, acceptor-acceptor) and hetero (donor-acceptor) labeled species. Only complexes with appropriate donor–acceptor stoichiometry were included in subsequent analyses.

Fluorescence lifetime (τ) analysis confirmed that the conjugated fluorophores retained their photophysical properties and sufficient rotational freedom for reliable FRET measurements. τ histograms of both conjugated and free fluorophores were fitted using a biexponential decay model, from which amplitude-weighted average lifetimes were calculated. For TmrABNBD, average τ values were 0.93 ± 0.02 ns (LD555) and 1.52 ± 0.01 ns (LD655), whereas TmrABPG exhibited average τ values of 0.95 ± 0.02 ns (LD555) and 1.65 ± 0.01 ns (LD655) (Fig. 1–Fig. S2a). By comparison, free dyes in buffer displayed lifetimes of 1.12 ± 0.02 ns (LD555) and 1.29 ± 0.01 ns (LD655). The similarity between free and conjugated dye lifetimes, together with comparable values across variants, indicates minimal protein-induced quenching and supports sufficient orientational averaging for quantitative smFRET analysis18,31–34.

Ensemble ATP titration (0–10 mM ATP) showed ATP-dependent donor quenching and acceptor sensitization (Fig. 1–Fig. S2b–d). ATP-induced fluorescence changes provided a quantitative readout of conformational transitions and enabled estimation of apparent equilibrium dissociation constants for ATP binding (Kd, ATP) in labeled TmrAB variants (Fig. 1–Fig. S2e–g). The measured apparent Kd, ATP values were 51 ± 38 μM for TmrABNBD, 68 ± 25 μM for TmrABPG, and 95 ± 26 μM for the slow-turnover variant TmrABPG_EQ, consistent with prior biochemical measurements (∼100 µM for TmrAE523QB)12. Labeling did not measurably perturb ATP binding.

Previous studies on spin-labeled TmrABNBD demonstrated transport activity comparable to wild-type TmrAB16, while AV simulations confirmed that fluorophores at these positions do not interfere with ATP- or substrate-binding sites. For labeled TmrABPG variant, no previous transport activity had been reported; we therefore performed transport assays (Fig. 1–Fig. S3) which confirmed that LD555/LD655 labeling at the PG does not impair function, as labeled and wild-type TmrAB showed indistinguishable substrate transport in single-liposome assays (Fig. 1-Fig. S3c).

ATP-induced conformational switching resolved by smFRET

TmrAB was immobilized on PEGylated surfaces using a conformation-independent, TmrB-specific nanobody (Nb9F10S63C)11. Previous studies confirmed that Nb9F10S63C does not perturb TmrAB transport11,14, consistent with single-liposome transport assays performed in this study (Fig. 1–Fig. S3d). For surface immobilization, Nb9F10S63C was conjugated to maleimide-PEG11-biotin, serving as an anchor to attach TmrAB to surface-immobilized streptavidin (Fig. 2a). Donor and acceptor photons were recorded by a total-internal reflection fluorescence (TIRF) microscope using alternating laser excitation (ALEX; NanoImager) at 40 °C (Fig. 2b). Single-molecule localization, fluorescence-trajectory extraction, and background correction were performed using NanoImager software, followed by DeepFRET-based machine-learning trace classification and corrections for donor leakage, direct acceptor excitation, and differences in detection-efficiency18,35,36 (Fig. 2–Fig. S1 and Fig. 2–Fig. S2). FRET efficiency (E) and stoichiometry (S) were calculated from the corrected fluorescence trajectories (see Methods, Eq. 1 and Eq. 2), and 2D E versus S diagrams were constructed to verify that only complexes with appropriate donor–acceptor stoichiometry (S ≈ 0.5) were included in the smFRET analysis (Fig. 2c).

ATP-induced conformational changes of TmrAB analyzed by smFRET.

a, Experimental setup. TmrABNBD (left) and TmrABPG (right) were labeled with LD555/LD655 and immobilized on PEGylated coverslips via a biotinylated conformation-independent, TmrB-specific nanobody (Nb9F10S63C)11,14. b, smFRET imaging was performed using total internal reflection fluorescence (TIRF) microscopy with alternating laser excitation (ALEX; donor: 532 nm; acceptor: 640 nm). Emission was collected in donor (498-620 nm) and acceptor (662-710 nm) channels. Donor emission upon donor excitation (fDexDem), sensitized acceptor emission upon donor excitation (fDexAem), and direct acceptor emission upon acceptor excitation (fAexAem) were used to calculate FRET efficiency (E) and stoichiometry (S). Data were analyzed using DeepFRET35. c, Representative TmrABNBD trace in ATP-bound state (top) and corresponding E/S plot (bottom). Donor emission upon donor excitation is shown in green (photon counts per frame), acceptor emission upon donor excitation in orange, FRET efficiency (E) in black, and stoichiometry (S) in grey. The initial increase in donor fluorescence reflects photophysical equilibration and instrumental stabilization and does not affect E or S, which are ratio-based and time-invariant18. d,e, Population analysis. FRET efficiency (E) histograms for (d) TmrABNBD and (e) TmrABPG are shown for apo (top) and ATP-bound (bottom; 3 mM ATP) conditions. Histograms were fitted with two Gaussian populations corresponding to the apo state (blue; defined from apo measurements) and ATP-bound state (orange; defined from saturating ATP fits). Dotted vertical lines indicate mean E values of each population, and relative fractions (Gaussian areas) are summarized schematically in each panel.

FRET efficiency (E) histograms revealed two dominant Gaussian populations corresponding to apo and ATP-bound states (Fig. 2d,e). Low-amplitude features were excluded from quantitative analysis because they were not reproducibly observed across replicates, lacked structural support from cryo-EM and DEER/PELDOR data, introduced poorly constrained fitting parameters, and increased the risk of overfitting. In the absence of ATP, only the apo population was observed, whereas addition of 3 mM ATP induced the appearance of a second ATP-bound population. For TmrABNBD, the apo and ATP-bound populations exhibited mean E values of 0.58 and 0.88 (ΔE = 0.30), respectively, with ∼77% of molecules occupying the ATP-bound state. For TmrABPG, mean E values were 0.97 (apo) and 0.86 (ATP-bound) (ΔE = 0.11), with ∼80% of molecules in the ATP-bound state.

Distance estimates using a Förster radius of R0 = 63.5 Å (ref.37) yielded apparent distances of 60.2 Å (apo) and 45.6 Å (ATP-bound) for TmrABNBD, and 35.6 Å (apo) and 46.9 Å (ATP-bound) for TmrABPG. For TmrABNBD, the experimentally derived ΔR of 14.6 Å closely agreed with the AV simulations. In contrast, the smaller ΔR of 11.4 Å observed for TmrABPG deviated from simulated values, suggesting that the ATP-bound population at this site represents a mixture of rapidly interconverting states rather than a single well-defined state.

ATP dependence of conformational equilibria

To assess the ATP sensitivity, we quantified conformational responses across a wide range of ATP concentrations, spanning well below the reported Kd, ATP (∼100 µM for TmrAE523QB)12 up to physiologically relevant levels (3 mM ATP). ATP concentrations up to 3 mM were selected to approximate near-physiological conditions commonly used in in-vitro studies, where millimolar ATP levels ensure saturation of nucleotide-dependent conformational transitions and facilitate comparison with previous biochemical and structural analyses10–12,14. smFRET measurements revealed dose-dependent population shifts: TmrABNBD shifted from a low-FRET apo state (E = 0.58) to a high-FRET ATP-bound state (E = 0.88), whereas TmrABPG shifted from a high-FRET apo state (E = 0.97) to a lower-FRET ATP-bound state (E = 0.86) (Fig. 3a,c). Langmuir isotherm fits yielded apparent Kd, ATP values of 13 ± 1 μM for TmrABNBD and 2 ± 1 μM for TmrABPG (Fig. 3b,d), significantly lower than Kd ATP determined in ensemble FRET measurements (Fig. 1–Fig. S2e–g). This difference likely arises from the difficulty of deconvoluting overlapping FRET populations at sub-Kd, ATP concentrations, particularly for TmrABPG, where state assignment is less well separated. Despite this quantitative offset, both approaches consistently indicate ATP saturation well below physiological concentrations and therefore support the same mechanistic conclusion that ATP binding drives conformational switching in TmrAB.

ATP-dependent shifts in smFRET populations of TmrAB.

a,c, Increasing ATP concentrations (0-3 mM) progressively shifted the population between apo and ATP-bound conformations for (a) TmrABNBD and (c) TmrABPG. FRET efficiency (E) histograms were fitted with two Gaussian populations corresponding to the ATP-free state (blue; defined from apo samples) and the ATP-bound state (orange; determined from saturating ATP conditions). Dotted vertical lines indicate the mean E values of each state, and relative fractions (Gaussian areas) are summarized schematically in each panel. b,d, ATP-binding curves obtained by plotting the fraction of molecules in the ATP-bound states as a function of ATP concentration for (b) TmrABNBD (reporting NBD dimerization) and (d) TmrABPG (reporting PG opening). Data were fitted with a Langmuir isotherm to determine the apparent dissociation constant Kd, ATP of each variant.

Trapping of TmrABPG reveals a previously hidden outward-facing open state

At 3 mM, ∼80% of TmrABPG complexes populated the ATP-bound state (E = 0.86), closely matching the ATP-bound population (∼77% NBD-dimerized) observed for TmrABNBD (Fig. 3). This agreement indicates that both labeling strategies consistently report ATP-dependent conformational changes. However, the ATP-induced shift observed for the periplasmic-gate reporter TmrABPG (ΔE = 0.11, ΔR = 11.4 Å; Fig. 3c) was substantially smaller than predicted by AV simulations (ΔE = 0.27, ΔR = 23.2 Å; Table 1). This discrepancy suggests that the ATP-bound population at E = 0.86 represents an unresolved ensemble, potentially comprising OFopen and OFoccluded conformations as well as the post-hydrolysis asymmetric unlocked return states (URasym and URasym*), which cannot be distinguished from OFoccluded within the current FRET geometry11.

Earlier version: To determine whether these states are kinetically unresolved by smFRET, we applied three complementary strategies to arrest the OFopen conformation of TmrABPG: (i) a slow-turnover mutant (TmrABPG_EQ), (ii) Mg2+ depletion using EDTA, and (iii) trans-inhibition by high concentrations of the substrate peptide R9L (RRYQKSTEL) (Fig. 4). Slow-turnover variants have previously enabled structural separation of OFopen and OFoccluded states11,12,14,15. Mg2+ depletion blocks ATP hydrolysis while preserving ATP binding, thereby enabling rapid and reversible trapping14. We further hypothesized that trans-inhibition by peptide binding sterically restricts PG closure and would therefore stabilize OFopen in a dose-dependent manner26,27.

Revised version: To determine whether these states are kinetically unresolved by smFRET, we applied three complementary strategies to arrest the OFopen conformation of TmrABPG: (i) a slow-turnover mutant (TmrABPG_EQ), (ii) Mg²⁺ depletion using EDTA, and (iii) trans-inhibition by high concentrations of the substrate peptide R9L (RRYQKSTEL) (Fig. 4). Slow-turnover variants have previously enabled structural separation of OFopen and OFoccluded states11,12,14,15. Mg²⁺ depletion blocks ATP hydrolysis while preserving ATP binding, thereby enabling rapid and reversible trapping14. We further hypothesized that trans-inhibition by peptide binding sterically restricts PG closure and would therefore stabilize OFopen in a dose-dependent manner26,27.

Identification of the outward-facing open (OFopen) conformation.

a–c, Three complementary approaches were used to resolve the OFopen state: (a) the slow-turnover variant TmrABPG_EQ, (b) imaging in Mg2+-free buffer supplemented with EDTA, and (c) stabilizing via trans-inhibition using high concentrations of peptide substrate R9L (0.3–2 mM). FRET efficiency (E) histograms were fitted with three Gaussian populations corresponding to the ATP-free state (blue), the ATP-bound state (orange), and the OFopen state (green). All three strategies reveal a distinct OFopen population. Dotted vertical lines indicate the mean E values of each state, and relative fractions (Gaussian areas) are summarized schematically in each panel. d, Comparison of inter-residue distances. Distances (Å) between selected residues on the NBDs and PG of TmrAB (Cβ–Cβ) were determined using smFRET (this study, detergent-solubilized TmrAB), cryo-EM structures of nanodisc-reconstituted TmrAB (PDB 6RAH, 6RAN)11; accessible-volume (AV) simulations (this study), and PELDOR/DEER measurements of detergent-solubilized TmrAB16.

In the absence of ATP, either without nucleotide or in the presence of ADP (3 mM), the slow-turnover variant TmrABPG_EQ populated a single high-FRET state (E = 0.97) (Fig. 4a, top and middle). These results indicate that ADP binding alone is insufficient to promote either NBD dimerization or PG opening, which is consistent with previous biochemical and structural observations11,12. Upon ATP addition (3 mM), however, the conformational landscape diverged sharply from that of TmrABPG. Instead of the two-state distribution observed for TmrABPG (apo: E = 0.97, ∼20%; ATP-bound: E = 0.86, ∼80%) (Fig. 2e, bottom), TmrABPG_EQ exhibited three well-resolved populations with mean E values of 0.97 (∼14%), E = 0.86 (∼55%), and E = 0.63 (∼31%) (Fig. 4a, bottom).

As in TmrABPG, the high-FRET population (E = 0.97) corresponds to the IF state, whereas the intermediate-FRET population (E = 0.86) likely represents a dynamic equilibrium between OFopen and OFoccluded conformations, as suggested by cryo-EM analyses11. Post-hydrolysis return states (URasym and URasym*) are expected to be minimally populated in TmrABPG_EQ because of its drastically reduced ATP hydrolysis rate12. Notably, the low-FRET ATP-bound population (E = 0.63) was entirely absent in wild-type TmrABPG. The transition from E = 0.97 to E = 0.63 corresponds to ΔE = 0.34 and ΔR = 22.5 Å, in close agreement with the IF→OFopen distance predicted by AV simulations (ΔR = 23.2 Å; Fig. 1 and Table 1), thereby supporting assignment of the E = 0.63 population to the OFopen conformation.

Earlier version: Mg2+ depletion independently reproduced this three-state landscape. In the absence of Mg2+, wild-type TmrABPG transitioned from a single apo population (without ATP; E = 0.97) (Fig. 4b, top) to three ATP-bound populations (3 mM ATP; E = 0.97, ∼13%; E = 0.86, ∼56%; E = 0.63, ∼31%) (Fig. 4b, middle), closely resembling those observed for the slow-turnover variant (Fig. 4a, bottom). Reintroducing Mg2+ abolished the E = 0.63 population and restored the wild-type TmrABPG two-state distribution (Fig. 4b, bottom). This reversibility confirmed that the ATP-bound OFopen state (E = 0.63) is selectively revealed only when ATP hydrolysis is prevented.

Revised version: Mg²⁺ depletion independently reproduced this three-state landscape. In the absence of Mg2+, wild-type TmrABPG transitioned from a single apo population (without ATP; E = 0.97) (Fig. 4b, top) to three ATP-bound populations (3 mM ATP; E = 0.97, ∼13%; E = 0.86, ∼56%; E = 0.63, ∼31%) (Fig. 4b, middle), closely resembling those observed for the slow-turnover variant (Fig. 4a, bottom). Reintroducing Mg²⁺ abolished the E = 0.63 population and restored the wild-type TmrABPG two-state distribution (Fig. 4b, bottom). This reversibility confirmed that the ATP-bound OFopen state (E = 0.63) is selectively revealed only when ATP hydrolysis is prevented.

Finally, we tested whether periplasmic substrate binding shifts the conformational equilibrium of wild-type TmrABPG toward OFopen. In the presence of ATP (3 mM), increasing concentrations of peptide substrate (0.3–2 mM R9L) progressively enriched the E = 0.63 population from ∼20% to ∼38% (Fig. 4c). This dose-dependent stabilization mirrors the trans-inhibition behavior reported for human TAP1/2 and reflects the upper substrate-loading capacity of the transporter26.

Distance changes derived from smFRET closely match AV simulations, cryo-EM structures (PDB 6RAH, 6RAN)11, and DEER/PELDOR measurements15 (Fig. 4d), together validating assignment of the E = 0.63 population as the OFopen conformation.

Kinetics and thermodynamics of the transport cycle

Earlier version: ALEX-smFRET data were acquired with an effective temporal resolution of 200 ms (100 ms per excitation channel). Shorter integration times compromised the signal-to-noise ratio and precluded reliable FRET determination. To quantify conformational dynamics, we applied Hidden Markov Modeling (HMM) using MASH-FRET38, classifying traces as either static (single FRET state) or dynamic (multiple states). Approximately 95% of traces in each condition were classified as static. The predominance of static traces indicates that most conformational transitions occur faster than the 200 ms observation window rather than reflecting conformational arrest.

Revised version: ALEX-smFRET data were acquired with an effective temporal resolution of 200 ms (100 ms per excitation channel). Shorter integration times compromised the signal-to-noise ratio and precluded reliable FRET determination. To quantify conformational dynamics, we applied Hidden Markov Modeling (HMM) using MASH-FRET38, classifying traces as either static (single FRET state) or dynamic (multiple states). Approximately 95% of traces in each condition were classified as static. The predominance of static traces is consistent with most conformational transitions occurring faster than the 200 ms observation window rather than reflecting long-lived conformational arrest.

Earlier version: Although individual transitions could not be directly resolved, population-based analysis (Fig. 3), combined with biochemical turnover measurements (Fig. 1–Fig. S1c), allowed estimation of ATP-bound dwell times. At saturating ATP conditions well above the apparent Kd, ATP (3 mM, 40 °C), wild-type TmrAB exhibited a catalytic turnover rate of kcat = 2.57 ± 0.38 s-1 (Fig. 1–Fig. S1c), corresponding to a full transport cycle time (τcycle) of 395 ± 55 ms. ATP-bound dwell times (τd) were derived from population ratios obtained from Gaussian fits of the FRET efficiency histograms (Fig. 3; see Methods, Eq. 3 and Eq. 4). These analyses yielded ATP-bound dwell times of 304 ± 43 ms for TmrABNBD (∼77% ATP-bound) and 316 ± 44 ms for TmrABPG (∼80% ATP-bound), with the remaining ATP-free intervals (∼20–23%) accounting for 91 ms and 79 ms of the cycle, respectively.

Revised version: Although individual transitions could not be directly resolved, population-based analysis (Fig. 3), combined with biochemical turnover measurements (Fig. 1–Fig. S1c), allowed estimation of ATP-bound dwell times. Under saturating ATP conditions well above the apparent Kd, ATP (3 mM, 40 °C), wild-type TmrAB exhibited a catalytic turnover rate of kcat = 2.57 ± 0.38 s-1 (Fig. 1–Fig. S1c), corresponding to a full transport cycle time (τcycle) of 395 ± 55 ms. ATP-bound dwell times (τd) were derived from population ratios obtained from Gaussian fits of the FRET efficiency histograms (Fig. 3; see Methods, Eq. 3 and Eq. 4). These analyses yielded ATP-bound dwell times of 304 ± 43 ms for TmrABNBD (∼77% ATP-bound) and 316 ± 44 ms for TmrABPG (∼80% ATP-bound), with the remaining ATP-free intervals (∼20–23%) accounting for 91 ms and 79 ms of the cycle, respectively.

Added in the revised version: The kinetic behavior of the ATP-bound ensemble is consistent with the thermodynamic properties previously determined for TmrAB. Although ATP-bound conformations accounted for most of the transport cycle (∼77–80% occupancy), individual OFopen and OFoccluded states could not be resolved within the ATP-bound period (∼310 ms), indicating rapid interconversion between these conformations. This observation agrees with previous thermodynamic analyses showing that ATP binding drives the IF-to-OF transition with an overall free-energy change close to zero (ΔG = 0.26 ± 0.74 kJ mol⁻¹), despite substantial compensating enthalpic and entropic contributions (ΔH = 30.14 ± 2.80 kJ mol⁻¹; TΔS = 28.26 kJ mol⁻¹) (ref.15). This entropy–enthalpy compensation places the conformational equilibrium near its midpoint, minimizing the energetic bias between ATP-bound conformations and thereby promoting rapid exchange between OFopen and OFoccluded states. Accordingly, the unresolved ATP-bound ensemble observed during steady-state turnover is consistent with a thermodynamic landscape that enables dynamic sampling of multiple outward-facing conformations.

Earlier version: Together, these measurements establish a quantitative, single-molecule description of conformational state occupancies and dwell times throughout the catalytic cycle of a heterodimeric ABC transporter under active turnover conditions, providing a quantitative single-molecule framework for describing the dynamic landscape of its translocation cycle.

Revised version: Together, these measurements establish a quantitative single-molecule description of both the kinetic progression and thermodynamic state occupancies throughout the catalytic cycle of a heterodimeric ABC transporter under active turnover conditions. By integrating population analysis with biochemical turnover measurements, this framework provides complementary insight into the dynamic and energetic landscape that governs the translocation cycle.

Discussion

smFRET has become an indispensable tool for dissecting conformational dynamics of membrane proteins, including receptors, ion channels, and transporters, by directly linking structural transitions to functional states. With the ABC transporter family, however, smFRET studies have largely focused on monomeric or homodimeric systems23–25, leaving asymmetric heterodimeric transporters comparatively underexplored. Here, we extend single-molecule analysis to the heterodimeric type IV ABC transporter TmrAB, revealing dynamic features of the transport cycle that are inaccessible to ensemble-averaged approaches.

By positioning FRET reporters at the nucleotide-binding domains (NBDs) and periplasmic gate (PG), we directly monitored ATP-dependent coupling between chemical energy input and global conformational rearrangements. Importantly, NBD dimerization is not uniquely coupled to a single outward-facing state but can lead to either an outward-facing open (OFopen) or occluded (OFoccluded) conformation11,12. This decoupling highlights the need to monitor both cytosolic and periplasmic regions to resolve the transport mechanism.

A limitation of the present study is the use of two FRET reporter pairs. While these positions were selected to probe key functional regions and validated by structural and spectroscopic benchmarks, additional labeling positions could in principle resolve further intermediates. However, increasing the number of states would also complicate quantitative analysis due to partially overlapping FRET populations, particularly for TmrABPG. The current design therefore balances mechanistic resolution with robust state assignment.

Labeling sites previously validated for PELDOR/DEER spectroscopy10,15,16 were adapted for smFRET and rigorously benchmarked using accessible-volume (AV) simulations30, fluorescence lifetime analysis, and ensemble FRET titrations. This validation is essential because fluorophores impose stricter steric and rotational constraints than spin labels32. Collectively, these controls demonstrate that fluorophore attachment preserves native-like conformational behavior and ATP responsiveness. Fluorescence lifetime measurements further indicated efficient energy transfer without evidence of substantial protein-fluorophore quenching or restricted mobility18.

Single-molecule measurements resolved two dominant FRET populations corresponding to apo and ATP-bound states for both TmrABNBD and TmrABPG. ATP titrations from sub-Kd, ATP to physiological concentrations (3 mM ATP) produced gradual, concentration-dependent redistribution between these states, demonstrating high sensitivity of both reporters. Apparent Kd, ATP values derived from smFRET (2–13 µM) were lower than those obtained from ensemble FRET (50–100 µM). Importantly, both approaches agree that TmrAB reaches conformational saturation well below physiological ATP concentrations and retains Michaelis–Menten behavior.

Earlier version: Using three independent strategies—slow-turnover catalysis, Mg2+ depletion, and substrate trans-inhibition—we resolved this previously hidden OFopen conformation. The observed distance changes agree with cryo-EM structures11, PELDOR/DEER data10,15,16, and AV-based predictions, indicating that detergent-solubilized TmrAB samples a largely native-like conformational landscape.

Revised version: Using three independent strategies—slow-turnover catalysis, Mg²⁺ depletion, and substrate trans-inhibition—we resolved this previously hidden OFopen conformation. The observed distance changes agree with cryo-EM structures11, PELDOR/DEER data10,15,16, and AV-based predictions, indicating that detergent-solubilized TmrAB samples a largely native-like conformational landscape.

Although detergent and lipid nanodisc environments yield similar global states, fluorophore behavior may still be influenced by membrane proximity in lipid systems20,39,40. AV simulations suggest that dyes near the periplasmic gate may experience steric constraints due to partial membrane overlap. While negligible under detergent conditions, this effect should be considered in future studies of membrane-reconstituted systems.

Substrate addition shifted the conformational equilibrium in a concentration-dependent manner, progressively stabilizing the OFopen state. This behavior is consistent with trans-inhibition observed in other ABC transporters, including TAP1/226 and ABCC127, and reflects finite substrate loading capacity26.

Quantitative deconvolution of FRET populations enabled estimation of conformational state occupancies under physiological ATP concentrations. During active turnover, TmrAB populates the IF state (∼20%), OFopen (∼25%), and OFoccluded/post-hydrolysis states (URasym and URasym*) (∼55%; Fig. 5). The current labeling scheme does not allow direct separation of OFoccluded from post-hydrolysis states due to fluorophore placement at the noncanonical NBS. However, post-hydrolysis contributions are expected to be reduced in the slow-turnover mutant TmrABPG_EQ. To our knowledge, this represents the first single-molecule quantification of conformational equilibria in a heterodimeric ABC transporter under catalytic conditions. While IF substructure resolution is lower than in cryo-EM classifications11,14, smFRET uniquely captures ATP-bound intermediates that interconvert too rapidly for structural resolution.

Conformational state distribution and catalytic cycle of TmrAB under active turnover.

Earlier version: Schematic of the TmrAB transport cycle summarizing major conformational states and their estimated population distributions under physiological ATP concentrations (3 mM, 40 °C). a, The inward-facing apo state (IFnarrow and IFwide; blue arc) accounts for ∼20% of molecules and is characterized by separated NBDs and a cytosol-accessible substrate-binding cavity. Substrate binding stabilizes the IFwide conformation11. ATP binding induces NBD dimerization and formation of the ATP-bound ensemble. b,c, Under substrate-bound turnover conditions, TmrAB proceeds via (b) OFoccluded and (c) OFopen states in which substrate release occurs. In steady-state turnover, the ATP-bound ensemble rapidly interconverts between OFoccluded and OFopen, accounting for ∼25% of the ATP-bound population (green circle). These transitions occur faster than the ∼200 ms temporal resolution of smFRET measurements, resulting in an averaged signal under turnover conditions. OFoccluded is proposed as an obligate intermediate between IF and OFopen, preventing substrate backflow by maintaining an occluded binding cavity during rearrangements of the PG and NBDs. Although a substrate-bound OFoccluded state has not been directly observed for TmrAB, its existence is supported by structures of the homodimeric type IV transporter BmrA47. Reduced ATP hydrolysis or substrate trans-inhibition enables trapping of the OFopen. state. d,e, ATP hydrolysis and phosphate (Pi) release generate post-hydrolysis return states (d) URasym and (e) URasym*. Subsequent ADP release restores the apo IF conformation, completing the transport cycle. Overall, the ATP-bound phase (b–e) represents ∼55% occupancy (orange arc) with an estimated dwell time of ∼310 ms, whereas the apo/ATP-rebinding phase (a) lasts ∼90 ms, yielding a total cycle time of ∼400 ms (kcat = 2.57 s−1). TmrA is shown in blue, TmrB in yellow, substrate as a green diamond, and nucleotides as orange symbols. Dotted grey boxes indicate the approximate position of the NBD dimer interface.

Revised version: Schematic of the TmrAB transport cycle summarizing major conformational states and their estimated population distributions under physiological ATP concentrations (3 mM, 40 °C). a, The inward-facing apo state (IFnarrow and IFwide; blue arc) accounts for ∼20% of molecules and is characterized by separated NBDs and a cytosol-accessible substrate-binding cavity. Substrate binding stabilizes the IFwide conformation11. ATP binding induces NBD dimerization and formation of the ATP-bound ensemble. b,c, Under substrate-bound turnover conditions, TmrAB proceeds via (b) OFoccluded and (c) OFopen states in which substrate release occurs. In steady-state turnover, the ATP-bound ensemble rapidly interconverts between OFoccluded and OFopen, accounting for ∼25% of the ATP-bound population (green circle). These transitions occur faster than the ∼200 ms temporal resolution of smFRET measurements, resulting in an averaged signal under turnover conditions. OFoccluded is proposed as an obligate intermediate between IF and OFopen, preventing substrate backflow by maintaining an occluded binding cavity during rearrangements of the PG and NBDs. Although a substrate-bound OFoccluded state has not been directly observed for TmrAB, its existence is supported by structures of the homodimeric type IV transporter BmrA47. Reduced ATP hydrolysis or substrate trans-inhibition enables trapping of the OFopen. state. d,e, ATP hydrolysis and phosphate (Pi) release generate post-hydrolysis return states (d) URasym and (e) URasym*. Subsequent ADP release restores the apo IF conformation, completing the transport cycle. Overall, the ATP-bound phase (b–e) represents ∼55% occupancy (orange arc) with an estimated dwell time of ∼310 ms, whereas the apo/ATP-rebinding phase (a) lasts ∼90 ms, yielding a total cycle time of ∼400 ms (kcat = 2.57 s⁻¹). TmrA is shown in blue, TmrB in yellow, substrate as a green diamond, and nucleotides as orange symbols. Dotted grey boxes indicate the approximate position of the NBD dimer interface.

All experiments were performed at 40 °C for consistency with prior biochemical work and fluorophore stability. At the physiological temperature of T. thermophilus (68 °C), absolute rates are expected to increase, although relative state distributions may be preserved if the free-energy landscape remains unchanged. Despite a substantial fraction of static trajectories, ATP-dependent shifts indicate that TmrAB operates near the temporal resolution of the measurements. Combining state occupancies with enzymatic turnover yields an ATP-bound dwell time of ∼300 ms, consistent with previous biochemical estimates9,12.

Emerging microsecond-resolution smFRET approaches could directly resolve short-lived intermediates41, and multi-color FRET strategies42,43 may enable simultaneous monitoring of NBDs, PG, and nucleotide sites. These advances would allow direct detection of the OFoccluded states and post-hydrolysis dynamics.

Overall, the consistency between reporter positions, ensemble measurements, and structural data support the conclusion that the observed conformational states reflect native TmrAB behavior. While certain intermediates remain kinetically unresolved under turnover conditions, they are strongly supported by independent trapping approaches, reinforcing the robustness of the mechanistic model. Additional conditions, including nucleotide-state asymmetry and substrate-only binding, represent promising directions for future work.

In summary, this work establishes smFRET as a powerful approach for mapping the dynamic landscape of asymmetric ABC transporters. By quantitatively linking ATP binding, conformational equilibria, and kinetics at the single-molecule level, we define how chemical energy is converted into directional transport in heterodimeric ABC systems.

Methods

Expression, purification, and labeling of TmrAB

His10-tagged TmrAB variants were expressed in E. coli BL21(DE3) (Thermo Fisher Scientific) as described previously12. Cells were grown in high-salt LB media (Carl Roth) supplemented with 100 μg ml-1 ampicillin (PAA Laboratories) at 37 °C. At an OD600 of 0.5, expression was induced with 1 mM isopropyl β-D-thiogalactopyranoside (IPTG; Carl Roth), and cultures were incubated for 3 h at 37 °C. Cells were harvested by centrifugation (4,500 × g, 4 °C, 15 min) and stored at –80 °C.

For purification, cell pellets were resuspended in lysis buffer (20 mM HEPES-NaOH pH 7.5, 300 mM NaCl, 50 µg ml-1 lysozyme, 0.2 mM phenylmethylsulfonyl fluoride (PMSF)) and lysed by sonication. Cell debris was removed by centrifugation (18,000 × g, 4 °C, 35 min), and membranes were collected by ultracentrifugation (100,000 × g, 4 °C, 30 min). Membranes were solubilized for 2 h at 4 °C in purification buffer (20 mM HEPES-NaOH pH 7.5, 300 mM NaCl) containing 20 mM n-dodecyl β-D-maltoside (β-DDM; Carl Roth). After ultracentrifugation (100,000 × g, 30 min, 4 °C), the supernatant was incubated with Ni-NTA agarose (Bio-Rad) for 1 h at 4 °C. The resin was washed with 20 column volumes of wash buffer (20 mM HEPES-NaOH pH 7.5, 300 mM NaCl, 1 mM β-DDM) containing 50 mM imidazole, and TmrAB was eluted with elution buffer (20 mM HEPES-NaOH pH 7.5, 300 mM NaCl, 1 mM β-DDM, 300 mM imidazole).

For fluorophore labeling, TmrAB variants were conjugated via maleimide chemistry using LD555 and LD655 (Lumidyne Technologies). Labeling was carried out at a 1:10:10 molar ratio of protein to each dye in elution buffer for 3 h at 4 °C. Excess dye was quenched with 2 mM β-mercaptoethanol (Sigma-Aldrich), and the labeled protein was buffer-exchanged into size-exclusion chromatography (SEC) buffer (20 mM HEPES-NaOH pH 7.5, 150 mM NaCl, 1 mM β-DDM) using a Zeba Spin Desalting Column (Thermo Fisher Scientific). Unreacted fluorophores were removed by SEC on a Superdex 200 Increase 10/300 GL column (Cytiva). Labeling efficiency was determined by analytical SEC (Superdex 200 Increase 3.2/300; Cytiva) by monitoring absorbance at 280, 555, and 655 nm. To preserve sample integrity for smFRET measurements, TmrAB was purified and labeled within a single day, stored on ice, and imaged over the following two days.

Time-correlated single-photon counting (TCSPC)

Fluorescence lifetime measurements were performed using a FluoTime 100 spectrometer (PicoQuant) equipped for time-correlated single-photon counting (TCSPC). Experiments were carried out on labeled TmrAB in SEC buffer (20 mM HEPES-NaOH pH 7.5, 150 mM NaCl, 1 mM β-DDM). LD555 and LD655 were excited at 510 nm and 610 nm, respectively. Emission was collected using a 620/60 nm bandpass filter for LD555 and a BG4 700 nm long-pass filter for LD655. Photon arrival times were accumulated until the TCSPC histogram reached a peak count of 50,000 photons. Fluorescence decay curves were analyzed by fitting mono- or bi-exponential decay models using FluoFit software (PicoQuant) and amplitude-weighted average of fluorescence lifetime was calculated.

Nanobody production and purification

The nanobody Nb9F10S63C was expressed and purified as described previously11. Briefly, Nb9F10S63C was produced in E. coli BL21(DE3) cells grown in Terrific Broth (TB; Carl Roth) supplemented with 100 μg ml-1 ampicillin at 37 °C. At an OD600 of 0.6, expression was induced with 1 mM IPTG, followed by overnight incubation at 28 °C. Cells were harvested by centrifugation (4,500 × g, 4 °C, 15 min) and stored at –80 °C. For purification, cell pellets were resuspended in nanobody lysis buffer (25 mM HEPES-NaOH pH 7.4, 300 mM NaCl, 15 mM imidazole, 0.5 mM PMSF) and disrupted by sonication. Cell debris was removed by centrifugation (18,000 × g, 4 °C, 35 min), and the clarified lysate was applied to Ni-NTA agarose equilibrated in potassium phosphate (KPi) buffer (25 mM KPi pH 6.5, 100 mM KCl, and 0.5 mM tris(2-carboxyethyl) phosphine (TCEP)). Bound nanobody was washed with 10 column volumes (CV) of KPi buffer and eluted with 8 CV of elution buffer (25 mM KPi pH 6.0, 20 mM KCl, 300 mM imidazole, 0.5 mM TCEP). Eluted fractions were pooled and further purified by cation exchange chromatography (CEX) on a HiTrap SP column (Cytiva) using a linear gradient from low-salt buffer (25 mM KPi pH 6.0, 20 mM KCl, 0.5 mM TCEP) to high-salt buffer (25 mM KPi pH 6.0, 500 mM KCl, 0.5 mM TCEP). The purified nanobody was concentrated and buffer-exchanged into nanobody SEC buffer (20 mM HEPES-NaOH pH 7.5, 150 mM NaCl) using Zeba spin desalting columns, followed by SEC (Superdex 200 Increase 10/300 GL; Cytiva). For site-specific conjugation, Nb9F10S63C was incubated with a biotin-PEG11-maleimide linker (Sigma-Aldrich) at a 1.2:1 molar ratio of protein to linker in the presence of 0.5 mM TCEP for 2 h at 4 °C. Excess linker was removed by desalting on Zeba Spin Desalting Columns, followed by a final SEC step (Superdex 200 Increase 10/300 GL).

SDS-PAGE

Earlier version: The purity of TmrAB samples was assessed by SDS-PAGE. Resolving gels (10%) were prepared using 10% (w/v) acrylamide, 0.5 M Tris-HCl (pH 8.8), 0.13% (w/v) SDS, 0.05% (w/v) ammonium persulphate (APS), and 0.25% (v/v) N,N,N’,N’-tetramethylethylenediamine (TEMED). Stacking gels contained 4.3% (w/v) acrylamide, 0.5 M Tris/HCl (pH 6.8), 0.09% (w/v) SDS, 0.09% (w/v) APS, and 0.33% (v/v) TEMED. Gels were used immediately or stored at 4 °C for up to 4 weeks. Protein samples were mixed with 4× SDS loading buffer containing dithiothreitol (DTT; Sigma-Aldrich) and heated at 90 °C for 5 min before loading. Electrophoresis was performed at a constant voltage of 120 V using 1× SDS running buffer (25 mM Tris-HCl pH 8.8, 192 mM glycine, 0.1% SDS). Proteins were visualized by staining with InstantBlue™ Protein Stain (Abcam) for 1 h at room temperature with gentle agitation and imaged using a Fusion FX system (Vilber).

Revised version: The purity of TmrAB samples was assessed by SDS-PAGE. Resolving gels (10%) were prepared using 10% (w/v) acrylamide, 0.5 M Tris-HCl (pH 8.8), 0.13% (w/v) SDS, 0.05% (w/v) ammonium persulphate (APS), and 0.25% (v/v) N,N,N’,N’-tetramethylethylenediamine (TEMED). Stacking gels contained 4.3% (w/v) acrylamide, 0.5 M Tris/HCl (pH 6.8), 0.09% (w/v) SDS, 0.09% (w/v) APS, and 0.33% (v/v) TEMED. Gels were used immediately or stored at 4 °C for up to 4 weeks. Protein samples were mixed with 4× SDS loading buffer containing dithiothreitol (DTT; Sigma-Aldrich) and heated at 90 °C for 5 min before loading. Electrophoresis was performed at a constant voltage of 120 V using 1× SDS running buffer (25 mM Tris-HCl pH 8.8, 192 mM glycine, 0.1% SDS). Proteins were visualized by staining with InstantBlueTM Protein Stain (Abcam) for 1 h at room temperature with gentle agitation and imaged using a Fusion FX system (Vilber).

ATPase activity assay

The ATPase activity of β-DDM-solubilized TmrABwt was quantified using a Malachite Green-based colorimetric assay as described previously44. Detergent-solubilized TmrAB (60 nM) was incubated in ATPase buffer (20 mM HEPES-NaOH pH 7.5, 150 mM NaCl, 2 mM MgCl2, 1 mM β-DDM) containing 3 mM ATP (Sigma-Aldrich) at 40 °C for 7 min. Autohydrolysis controls were prepared by incubation of ATP in ATPase buffer without protein. Reactions were quenched by adding 20 mM H2SO4, followed by incubation with 3 mM Malachite Green (Thermo Fisher Scientific), 0.2% (v/v) Tween20 (Carl Roth), and 1.5% (w/v) ammonium molybdate (Carl Roth) for 10 min at room temperature. The absorbance at 620 nm was recorded on a CLARIOstar v.5.20 R5 plate reader (BMG LABTECH).

Ensemble FRET measurements

The ATP binding and FRET characteristics of selected TmrAB variants were assessed by ensemble FRET. Labeled TmrAB (100 nM) was incubated with increasing concentrations of ATP at 42 °C for 5 min. Donor-excited emission was recorded from 550–700 nm with an excitation wavelength of 520 nm using a CLARIOstar v.5.20 R5 plate reader (BMG LABTECH). Acceptor emission intensities at 675 nm were plotted against ATP concentration and fitted with a hyperbolic function to determine the apparent dissociation constant (Kd, ATP) for each variant.

Reconstitution of TmrAB

TmrAB variants were reconstituted into proteoliposomes composed of E. coli polar lipids and dioleoylphosphatidylcholine (DOPC) at a molar ratio of 7:3. A protein-to-lipid ratio of 1:20 (w/w) was used, corresponding to approximately 50 TmrAB complexes per liposome. Extruded large unilamellar vesicles (LUVs) were destabilized by incubation with 0.3% (v/v) Triton X-100 for 30 min. TmrAB was then added to the detergent-destabilized liposomes and incubated for 30 min at 8 °C under gentle rotation. Detergent was removed gradually using sequential addition of polystyrene Bio-Beads SM-2 (Bio-Rad): two incubations at 40 mg ml-1 (1 h and overnight), followed by two incubations at 80 mg ml-1 (1 h each). Proteoliposomes were subsequently harvested by ultracentrifugation (100,000 × g, 30 min, 4 °C) and resuspended in the appropriate buffer.

Single liposome transport assay

Single liposome transport assay was performed as described previously12,13. Wild-type TmrAB or LD555/LD655-labeled TmrABPG (0.6 µM each), reconstituted into liposomes (∼50 TmrAB per liposome), was mixed with fluorescently labeled peptide RRYCFKSTEL (30 µM C4F; F, fluorescein), 3 mM ATP or ADP, and 5 mM MgCl2 in liposome buffer (25 mM HEPES-NaOH pH 7.5, 150 mM NaCl, 5% (v/v) glycerol). Where indicated, nanobody Nb9F10S63C (0.6 µM) was added. Samples were incubated for 5 min at 40 °C (unless specified otherwise), followed by addition of ice-cold ethylenediaminetetraacetic acid (EDTA; 10 mM). Proteoliposomes were then washed twice by centrifugation (270,000 x g, 30 min) to remove non-transported peptide. Mean fluorescence intensities (MFIs) of translocated substrates were quantified by flow cytometry using a FACSMelody (Waters Biosciences) and analyzed in FlowJo 10.6.1 (Waters Biosciences). The gating strategy used to determine fluorescence intensities of single liposomes from raw flow cytometry data consisted of three steps: First, forward scatter area versus side scatter area (FSC-A vs SSC-A) was used to separate populations based on size and granularity. Second, forward scatter height versus forward scatter area (FSC-H vs FSC-A) was used to exclude aggregates and retain singlet events. Finally, fluorescence intensities of the gated singlet population were quantified at the wavelength of interest (527 nm), and the resulting values were used to calculate mean fluorescence intensity (MFI) of the proteoliposome population.

Solid-phase peptide synthesis

The peptides used in this study were synthesized and labeled as described previously14,45. The 9-mer peptides RRYQKSTEL (R9L) and RRYCFKSTEL (C4F; F, fluorescein) were synthesized on preloaded Fmoc-L-Leu resin using a Liberty microwave-assisted peptide synthesizer (CEM) according to a standard protocol (54 W, 3 min, 75 °C). Each coupling step was performed twice using 0.2 M fluorenylmethoxycarbonyl (Fmoc)-protected amino acids, 0.5 M O-benzotriazole-N,N,N’,N’-tetramethyluronium hexafluoro-phosphate (HBTU), and 1-hydroxybenzotriazole hydrate (HOBt H2O). Fmoc deprotection was carried out with 20% (v/v) piperidine and 0.1 M HOBt H2O in dimethylformamide (DMF). Peptides were cleaved from the resin using a cleavage cocktail containing 92.5% (v/v) trifluoroacetic acid (TFA), 2.5% (v/v) H2O, 4.5% (v/v) thioanisole, and 0.5% (v/v) 1,2-ethanedithiol (EDT) for 1.5 h at room temperature. Cleaved peptides were precipitated in ice-cold diethyl ether (Et2O), pelleted, dissolved in tert-butanol (tBuOH)/water (4:1, v/v), and lyophilized. Peptides were purified by reverse-phase (RP) C18 HPLC (PerfectSil C18 column; MZ-Analysentechnik) using buffer A (Milli-Q water containing 0.05% (v/v) TFA) and buffer B (acetonitrile containing 0.05% (v/v) TFA). Peptide identity and purity were confirmed by LC-MS. For fluorescence labeling, the peptide RRYCFKSTEL was dissolved in 33% (v/v) DMF in PBS (10 mM Na2HPO4, 1.8 mM KH2PO4, 137 mM NaCl, 2.7 mM KCl, pH 7.4) and incubated for 1 h at room temperature with a 1.3-fold molar excess of 5-iodoacetamido fluorescein (Merck).

Functionalization of glass slides for single-molecule FRET analysis

Glass coverslips used for TmrAB immobilization in smFRET experiments were functionalized by PEGylation as described previously46. Coverslips (Carl Roth) were cleaned by sequential sonication in Milli-Q water and analytical-grade acetone (>99.9%; VWR International), followed by oxygen plasma treatment (0.3 mbar, 80% power, 15 min) using a Zepto plasma cleaner (Diener) and a 10 min incubation in methanol (Avantor, Gliwice, PL). Coverslips were then silanized by incubation for 30 min in a solution of 100 ml methanol, 5 ml acetic acid, and 3 ml 3-aminopropyltrimethoxysilane (APTES; Tokyo Chemical Industry). After silanization, slides were rinsed four times with methanol and dried under a nitrogen stream. Surface functionalization was achieved using a mixture of biotinylated-PEG (4 mol%) and nonbiotinylated-PEG (96 mol%; Rapp Polymer). The PEG solution was sandwiched between two coverslips and incubated overnight in a humidity chamber. Coverslips were then rinsed thoroughly with Milli-Q water and dried under nitrogen. To enhance passivation, a second PEGylation step was performed using 25 mM CH3-PEG-NHS (333 Da; Thermo Fisher Scientific) under the same conditions. Finally, slides were rinsed with Milli-Q water, dried under nitrogen, and stored at –20 °C under argon until use.

Single-molecule FRET imaging

Single-molecule FRET (smFRET) experiments were performed using a flow chamber system (Ibidi). Chambers were assembled by placing a biotin-PEG-functionalized glass slide onto a μ-Slide I Luer Family flow channel (Ibidi), with the functionalized surface facing inward. All buffers and Milli-Q water were filtered through 0.2 μm filters (Sigma-Aldrich). Chambers were washed with 1 ml Milli-Q water and incubated with 0.2 mg ml-1 streptavidin (Sigma-Aldrich) at 4 °C for 30 min to allow binding to the biotin-PEG surface. Unbound streptavidin was removed by washing with 1 ml Milli-Q water. The surface was then treated with 0.3 mg ml-1 biotinylated-PEG11-Nb9F10S63C at 4 °C for 45 min, followed by flushing with 2 ml SEC buffer (20 mM HEPES-NaOH pH 7.5, 150 mM NaCl, 1 mM β-DDM). Detergent-solubilized, fluorophore-labeled TmrAB (100 nM) was added and incubated at 4 °C for 1 h. Unbound protein was removed by five washes with 1 ml TmrAB-SEC buffer. Chambers were equilibrated with 1 ml of imaging buffer containing 25 mM HEPES-NaOH (pH 7.5), 150 mM NaCl, 3 mM MgCl2, 50 mM glucose, 5 mM Trolox, 7.5 U ml-1 pyranose oxidase, and 1 kU ml-1 catalase, supplemented with the desired ATP concentration. For EDTA trapping experiments, MgCl2 was omitted and replaced with 3 mM ethylenediaminetetraacetic acid (EDTA; Sigma-Aldrich). smFRET data were acquired at 40 °C using alternating laser excitation (ALEX) on a total internal reflection fluorescence (TIRF) microscope (NanoImager S, ONI, Oxford, UK). Typically, 600 frames were recorded per region of interest (ROI) with 100 ms exposure time. Laser powers were 0.8 mW cm-2 (532 nm) and 0.9 mW cm-2 (640 nm). Data were recorded in 1-min intervals, except for TmrABNBD variant in apo and 3 mM ATP conditions, where 3-min intervals were used to confirm that conformational transitions do not occur on timescales longer than one minute due to the reduced temperature.

Single-molecule FRET data analysis

Earlier version: Single-molecule FRET (smFRET) measurements were performed using alternating laser excitation (ALEX), allowing assignment of detected photons based on both excitation and emission wavelengths. Photon counts for each molecule were extracted using NanoImager software (ONI NanoImager, Development Build) and classified into three detection channels: donor excitation with donor emission (fDexDem), donor excitation with acceptor emission (fDexAem), and acceptor excitation with acceptor emission (fAexAem). Traces were analyzed with DeepFRET35 and manually curated. To minimize bias, a second researcher independently curated traces from both ATP-free and 3 mM ATP samples, yielding 98% overlap between curations. FRET efficiency (E) and stoichiometry (S) were calculated as:

Revised version: Single-molecule FRET (smFRET) measurements were performed using alternating laser excitation (ALEX), allowing assignment of detected photons based on both excitation and emission wavelengths. Photon counts for each molecule were extracted using NanoImager software (ONI NanoImager, Development Build) and classified into three detection channels: donor excitation with donor emission (fDexDem), donor excitation with acceptor emission (fDexAem), and acceptor excitation with acceptor emission (fAexAem). Traces were analyzed with DeepFRET35 and manually curated. To minimize bias, a second researcher independently curated traces from both ATP-free and 3 mM ATP samples, yielding 98% overlap between curations. FRET efficiency (E) and stoichiometry (S) were calculated as: (Eq. 1) E=fDexAemfDexDem+fAexAem

Removed in the revised version: (1) E=fDexAemfDexDem+fAexAem

Removed in the revised version: (2) S=fDexDem+fDexAemfDexDem+fDexAem+fAexAem

Population analysis was performed by constructing one-dimensional histograms of FRET efficiency (E) and stoichiometry (S) using OriginPro 2024 (OriginLab). Histograms were fitted with two Gaussian distributions corresponding to the ATP-free state (defined from apo samples) and the ATP-bound state (defined by a two-component fit at saturating ATP). Hidden Markov Modeling (HMM) of individual traces was performed using MASH-FRET38 to distinguish dynamic from static molecules within each sample.

ATP-bound dwell times and distribution of conformational states

Removed in the revised version: The ATP-bound dwell time (τd) was estimated as:

Earlier version: (3) τd=τcycle×fATP-bound where fATP-bound is the fraction of ATP-bound molecules derived from Gaussian fits of the FRET efficiency histograms (Fig. 3). This approach assumes (i) ATP hydrolysis occurs exclusively at the canonical NBS, with negligible contribution from the noncanonical site9,11,15, and (ii) the majority of molecules are catalytically competent and continuously cycling.

Revised version: The ATP-bound dwell time (τd) was estimated as: (Eq. 3) τd=τcycle×fATP−bound where fATP-bound is the fraction of ATP-bound molecules derived from Gaussian fits of the FRET efficiency histograms (Fig. 3). This approach assumes (i) ATP hydrolysis occurs exclusively at the canonical NBS, with negligible contribution from the noncanonical site9,11,15, and (ii) the majority of molecules are catalytically competent and continuously cycling.

Removed in the revised version: The distribution of conformational states within ATP-bound FRET population (E = 0.86) of TmrABPG under turnover conditions (3 mM ATP) was determined using a two-state model:

Earlier version: (4) E=fopen×Eopen+fother×Eother where fopen is the fraction of OFopen state (Eopen = 0.63) readily resolved under trapping conditions (Fig. 4), and fother represents the combined fraction of PG-closed states (OFoccluded/URasym/URasym*, Eother = 0.97).

Revised version: The distribution of conformational states within ATP-bound FRET population (E = 0.86) of TmrABPG under turnover conditions (3 mM ATP) was determined using a two-state model: (Eq. 4) E=fopen×Eopen+fother×Eother where fopen is the fraction of OFopen state (Eopen = 0.63) readily resolved under trapping conditions (Fig. 4), and fother represents the combined fraction of PG-closed states (OFoccluded/URasym/URasym*, Eother = 0.97).